# Activity 1: Implementing Insertion Sort

Implement the insertion sort algorithm:

INSERTION-SORT(A,n)
```python showLineNumbers
for i = 2 to n
	key  = A[i]
    //Insert A[i] into the sorted subarray A[1 : i-1]
    j = i - 1
    while j > 0 and A[j] > key
    	A[j + 1] = A[j]
        j = j - 1
    A[j + 1] = key

```


<!-- grade_id: 22367d, points: 2, grade: true, solution: true -->
```python
def insertion_sort(A, n):
    '''takes integer list A[1 : n] and size n of A as a parameters, A is the list to sort, A IS modified'''
    
    for i in range (2, n+1):
        key  = A[i]
        j = i - 1
        
        while j > 0 and A[j] > key:
            A[j + 1] = A[j]
            j = j - 1
        A[j + 1] = key
```

<!-- grade_id: 8f8172, points: 2, grade: true, solution: false -->
```python
# Unit test for insertion sort
# These tests will not check if your function implements the GIVEN algorithm. You need to DO THAT YOURSELF!
from nose.tools import assert_equal
z = ["empty"]
insertion_sort(z,0)
assert_equal(z, z)
a = ["empty", 1]
insertion_sort(a,1)
assert_equal(a, a)
b = ["empty", 5,4,33,2]
insertion_sort(b,4)
assert_equal(b, ['empty',2,4,5,33])
c = ["empty",5,2,4,6,1,3,4,2,22,1]
insertion_sort(c, 10)
assert_equal(c, ['empty', 1, 1, 2, 2, 3, 4, 4, 5, 6, 22])

```