# Activity 12: Building a Max Heap

## Exercise 1

In the below cell, implement the PARENT, LEFT, and RIGHT  from Activity 11 First Heaps.

<!-- grade_id: ecf69e, points: 2, grade: true, solution: true -->
```python
#Exercise 1
def parent(i):
    '''return the index of the parent of node i'''
    return i//2

#Exercise 1
def left(i):
    '''return the index of the left child of node i'''
    return 2*i
    
#Exercise 1    
def right(i):
    '''return the index of the right child of node i'''
    return 2*i + 1

#Exercise 2
def max_heapify(A,i):
    '''Takes a list A[1..n] of size n and an index i into the list
    MAX-HEAPIFY assumes that the binary trees rooted at LEFT[i] RIGHT[i] are max-heaps, 
    but that A[i] might be smaller than its children, thus violating the max-heap property.
    MAX-HEAPIFY lets the value at A[i] “ﬂoat down” in the max-heap so that the subtree 
    rooted at index i obeys the max-heap property
    A is modified'''
    
    l = left(i)
    r = right(i)
    
    if l <= (len(A)-1) and A[l] > A[i]:
        largest = l
    else:
        largest = i
        
    if r <= (len(A)-1) and A[r] > A[largest]:
        largest = r
        
    if largest != i:
        temp = A[i]
        A[i] = A[largest]
        A[largest] = temp
        
        max_heapify(A, largest)

#Exercise 3
def build_max_heap(A,n):
    '''Takes an unordered list A[1..n] of size n and produces a max-heap
    A is modified'''
    
    for i in range (n//2, 0, -1):
        max_heapify(A, i)
    
```

<!-- grade_id: 6d4fcf, points: 1, grade: true, solution: false -->
```python
# Unit tests for parent, left, and right
# These tests will not check if your functions implement the GIVEN algorithm. You need to DO THAT YOURSELF!
from nose.tools import assert_equal


b = ["empty",16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
assert_equal(b[left(1)], 4)
assert_equal(b[right(4)], 8)
assert_equal(b[parent(7)], 10)

```

## Exercise 2

Above, implement the MAX-HEAPIFY algorithm from Activity 11 First Heaps.

<!-- grade_id: f589cf, points: 1, grade: true, solution: false -->
```python
# Unit tests for max_heapify
# These tests will not check if your function implements the GIVEN algorithm. You need to DO THAT YOURSELF!
from nose.tools import assert_equal
z = ["empty"]
max_heapify(z,0)
assert_equal(z, z)

a = ["empty", 1]
max_heapify(a,1)
assert_equal(a, a)

max_heapify(b,2)
assert_equal(b, ['empty', 16, 14, 10, 8, 7, 9, 3, 2, 4, 1])


```

## Exercise 3

Above, implement the BUILD-MAX-HEAP algorithms from Activity 12 Building a Heap.

<!-- grade_id: cbf8d7, points: 1, grade: true, solution: false -->
```python
# Unit tests for build_max_heap
# These tests will not check if your function implements the GIVEN algorithm. You need to DO THAT YOURSELF!
from nose.tools import assert_equal
z = ["empty"]
build_max_heap(z, len(z)-1)
assert_equal(z, z)

a = ["empty", 1]
build_max_heap(a, len(a)-1)
assert_equal(a, a)

c = ["empty", 4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(c, len(c)-1)
assert_equal(c, b)

d = ["empty", 5, 3, 17, 10, 84, 19, 6, 22, 9]
build_max_heap(d, len(d)-1)
assert_equal(d, ['empty', 84, 22, 19, 10, 3, 17, 6, 5, 9])

c = ["empty", 5, 3, 17, 10, 84, 19, 6, 22, 9]
```