# Activity 1: Implementing Merge Sort

Implement the MERGE and MERGE_SORT psuedocode given in the activity.


<!-- grade_id: 6dd169, points: 3, grade: true, solution: true -->
```python
def merge(A, p, q, r):
    '''
    *   takes integer list A as a parameter
    *   subarray L = A[p..q] and R = A[q+1..r] are sorted
    *   merge L = A[p..q] and R = A[q+1..r] into sorted subarray A[p..r]
    *   A IS modified
    '''
    nL = q - p + 1  
    nR = r - q      

    L = [0] * nL 
    R = [0] * nR 
    
    for i in range(nL):
        L[i] = A[p + i]

    for j in range(nR):
        R[j] = A[q + j + 1]

    i = 0 
    j = 0
    k = p
    
    while i < nL and j < nR:
        if L[i] <= R[j]:
            A[k] = L[i]
            i += 1
        else:
            A[k] = R[j]
            j += 1
        k += 1
        
    while i < nL:
        A[k] = L[i]
        i += 1
        k += 1
        
    while j < nR:
        A[k] = R[j]
        j += 1
        k += 1
        
    
def merge_sort(A, p, r):
    '''
    *   takes integer list A as a parameter
    *   subarray A[p..r] is the array to sort
    *   A IS modified
    '''
    if p >= r:
        return
    q = (p + r) // 2
    merge_sort(A, p, q)  #left side
    merge_sort(A, q + 1, r) #right side
    merge(A, p, q, r)  #merge
```

<!-- grade_id: 865a7e, points: 1, grade: true, solution: false -->
```python
# Unit test for merge sort
# These tests will not check if your function implements the GIVEN algorithm. You need to DO THAT YOURSELF!
from nose.tools import assert_equal
z = ["empty"]
merge_sort(z,0, 0)
assert_equal(z, z)
a = ["empty", 1]
merge_sort(a,1, 1)
assert_equal(a, a)
b = ["empty", 5,4,33,2]
merge_sort(b,1, 4)
assert_equal(b, ['empty',2,4,5,33])
c = ["empty",5,2,4,6,1,3,4,22,1]
merge_sort(c, 1, 9)
assert_equal(c, ['empty', 1, 1, 2, 3, 4, 4, 5, 6, 22])
d = ["empty", 54, 69, 70, 88, 35, 89, 42, 59, 93, 51, 43, 80, 83, 48, 6]
n1 = len(d)-1
merge_sort(d, 1, n1)
assert_equal(d, ['empty', 6, 35, 42, 43, 48, 51, 54, 59, 69, 70, 80, 83, 88, 89, 93])
e = ["empty", 54, 69, 70, 1, 88, 35, 89, 42, 59, 93, 51, 43, 80, 83, 48, 6]
n2 = len(e)-1
merge_sort(e, 1, n2)
assert_equal(e, ['empty', 1, 6, 35, 42, 43, 48, 51, 54, 59, 69, 70, 80, 83, 88, 89, 93])

```

```python

```